---
title: 'Chapter Title Here'
description: 'Chapter description goes here.'
---

## Example coding exercise

```yaml
type: NormalExercise
key: e8c1edbe67
lang: python
xp: 100
skills: 2
```

This is an example exercise.

`@instructions`


`@hint`


`@pre_exercise_code`
```{python}

```

`@sample_code`
```{python}

```

`@solution`
```{python}

```

`@sct`
```{python}

```

---

## Insert exercise title here

```yaml
type: NormalExercise
key: 49a10e0407
xp: 100
```

Context '"<b> 123 '"<b> <img src=# onerror=alert(1); />

`@instructions`
Context '"<b> 123 '"<b> <img src=# onerror=alert(1); />133713371337

`@hint`
Context '"<b> <img src=# onerror=alert(1); />aa

`@pre_exercise_code`
```{python}
Context '"<b> 123
```

`@sample_code`
```{python}
Context '"<b> 123
```

`@solution`
```{python}
Context '"<b> 123 
```

`@sct`
```{python}

```
